# IST 718: Big Data Analytics

- Professor: Daniel Acuna <deacuna@syr.edu>
- TA: Tong Zeng <tozeng@syr.edu>

# Homework 1

## General instructions

Before downloading and submitting your work through Blackboard, remember to save and press `Validate` (or go to `Kernel`$\rightarrow$`Restart and Run All`). The tests shown in some cells might not be the only ones used by the professor.

# Part 2: Basic Python programming

## Question 2.1: Basic data types (5 pts)

Create variables containing specific types
- `a` integer type (`int`)
- `b` floating type (`float`)
- `c` string type (`str`)
- `d` list type (`list`)
- `e` dictionary type (`dict`)
- `f` tuple type (`tuple`)
- `g` set type(`set`)

For example,
```python
a = 1
```

<!-- grade_id: cell-44734e6d8cfc336f, points: 0, grade: false, solution: true -->
```python
# example
a = 1
b = 1.1
c = "abc"
d = [1,2,3]
e = {1:1,2:"Two","Three":3}
f = ("Data","Science")
g = {"Data","Science"}
# YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: cell-14b68ec3bb8458cb, points: 5, grade: true, solution: false -->
```python
assert type(a) == int
assert type(b) == float
assert type(c) == str
assert type(d) == list
assert type(e) == dict
assert type(f) == tuple
assert type(g) == set
```

## Question 2.2: Finding an element in a list (10 pts)

Using a for loop, create a function that receives a list `L` of integers and an element `e`. The function should return the index of the element `e` in `L`. Make use of the length of the list (e.g., `len(L)`) and `range` to access each of index of the list. 

- If there are no elements that contain `e` or L is empty,  return `-1`.
- Only return the first index that contains the element.

For example, the index of `2` in `[1,2,3]` is 1. Therefore, `find_element([1,2,3], 2)` should return 1

<!-- grade_id: cell-fc5a397c1d4bc38c, points: 0, grade: false, solution: true -->
```python
def find_element(L, e):
    """Find index of first ocurrence of e in list L. If not found, it returns -1"""
    # YOUR CODE HERE
    index = 0
    for i in L:
        if i == e:
            break
        index+=1    
    if index == len(L):
        index = -1
    return index
    raise NotImplementedError()
```

<!-- grade_id: cell-bd1c7805ca9d2fba, points: 10, grade: true, solution: false -->
```python
assert find_element([], 0) == -1
assert find_element([1], 0) == -1
assert find_element([1,2,3,4,5], 2) == 1
assert find_element([1,2,3,4,5], 5) == 4
assert find_element([1,2,2,2,1], 2) == 1
```

## Question 2.3: Finding elements 2 (5 pts)

Modify the code of Question 2.2 to return the index of the element only if it is found in odd indices (i.e., 1, 3, 5, etc). Hint: You can easily achieve this by looking at how `range` works (https://docs.python.org/3/library/stdtypes.html#typesseq-range)

<!-- grade_id: cell-9b7fc5e5d436ab8e, points: 0, grade: false, solution: true -->
```python
def find_element_odd(L, e):
    """Find index of first ocurrence of e in list L. If not found, it returns -1"""
    # YOUR CODE HERE
    index = 0
    for i in L:
        if i == e:
            break
        index+=1    
    if index == len(L) or index%2 == 0:
        index = -1
    return index
    raise NotImplementedError()
```

<!-- grade_id: cell-56f7074ffffef0c6, points: 5, grade: true, solution: false -->
```python
assert find_element_odd([], 0) == -1
assert find_element_odd([1], 0) == -1
assert find_element_odd([1,2,3,4,5], 2) == 1
assert find_element_odd([1,2,3,4,5], 5) == -1
assert find_element_odd([1,2,2,2,1], 2) == 1
```

## Question 2.3: Mean (10 pts)

There are several built-in functions for Python's lists of numbers. One of them is `sum` which returns the sum of all elements in a list. Use this function to create a function that computes the mean of a list of numberes. If the list is empty, return 0.

**Only use vanilla list not `numpy`**

<!-- grade_id: cell-313f91e6e47a866e, points: 0, grade: false, solution: true -->
```python
def mean(L):
    """Return mean of numerical elements in L. return 0 if L is empty"""
    if sum(L) != 0:
        return sum(L)/len(L)
    else:
        return 0
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e38588dabb73c551, points: 10, grade: true, solution: false -->
```python
assert mean([1]) == 1
assert mean([0,10]) == 5
assert mean([-10.5, 0.5]) == -5
assert mean([]) == 0
```

## Question 2.4: Variance (20 pts)

Using only for loops, variable assignments, and reusing the function in Question 2.3, implement a function that computes the population variance of a list of numbers `L`. This is for a vector $x$ of length $n$, the population variance is

$\text{var}(x) = \frac{1}{n} \sum_{i=1}^n (x_i - \bar{x})^2$

where $\bar{x}$ is the mean of $x$

If the list is empty, return `float("inf")`, which represents infinity

**Only use vanilla list not `numpy`**

<!-- grade_id: cell-207394a94ccf1b12, points: 0, grade: false, solution: true -->
```python
def var(L):
    if len(L) == 0:
        return float("inf")
    else:
        return sum((i-mean(L))**2 for i in L)/len(L)
        
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# try it yourself
var([1])
```

Output:
```
0.0
```

```python
var([])
```

Output:
```
inf
```

<!-- grade_id: cell-e7c3507db876deba, points: 20, grade: true, solution: false -->
```python
"""Test variance function"""
assert var([]) == float("inf")
assert var([-10, 10]) == 100
assert var([1]) == 0
assert var([-1]) == 0
```

## Question 2.5: BuzzFizz (25 pts)

Remember the FizzBuzz problem from class. We will run a modification called BuzzFizz. By default, we will return the __squares__ of the numbers from 1 to 30. This time, instead of _printing_ the result, you will return a list of  with the results. But for numbers that are multiple of 3, we will produce the string "Buzz", and for numbers that are multiple of 4 we will produce the string "Fizz". If a number is multiple of 3 and 4, we will produce the string "BuzzFizz". 

The first 12 elements of the list should be

- 1
- 4
- "Buzz"
- "Fizz"
- 25
- "Buzz"
- 49
- "Fizz"
- "Buzz"
- 100
- 121
- "BuzzFizz"

<!-- grade_id: cell-2a659b9c25270e2d, points: 0, grade: false, solution: true -->
```python
def buzzfizz():
    results = []
    for i in range(1, 31):
    # YOUR CODE HERE
        if i%3==0 and i%4==0:
            results.append("BuzzFizz")
        elif i%4==0:
            results.append("Fizz")
        elif i%3==0:
            results.append("Buzz")
        else:
            results.append(i**2)      
    # raise NotImplementedError()
    return results
```

<!-- grade_id: cell-5748318ee26f4e7a, points: 25, grade: true, solution: false -->
```python
assert len(buzzfizz()) == 30
assert buzzfizz()[0:6] == [1, 4, 'Buzz', 'Fizz', 25, 'Buzz']
assert buzzfizz()[23] == "BuzzFizz"
```

```python

```